Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Vaishnavi Vukku"
COLLABORATORS = "Professor Zia"
```

---

<!-- grade_id: cell-63c5a6a20b2527de, points: 0, grade: false, solution: true -->
```python
# 4 points
def practice_1(name):
    # Complete this function to use the `name` input to print the following:
    # Hello, <name>, nice to meet you!
    # Example: name = Brian
    # Print: Hello, Brian, nice to meet you!
    
    # YOUR CODE HERE
    print("Hello, {}, nice to meet you!".format(name))
```

<!-- grade_id: cell-309b995cf2723861, points: 4, grade: true, solution: false -->
```python
from unittest.mock import patch
with patch('__main__.print') as mock_print:
    practice_1('Brian')
mock_print.assert_called_once_with('Hello, Brian, nice to meet you!')

with patch('__main__.print') as mock_print:
    practice_1('John')
mock_print.assert_called_once_with('Hello, John, nice to meet you!')

with patch('__main__.print') as mock_print:
    practice_1('Jane')
mock_print.assert_called_once_with('Hello, Jane, nice to meet you!')
```

<!-- grade_id: cell-d4cd235afa6ba9c9, points: 0, grade: false, solution: true -->
```python
# 4 points
def practice_2(x, y, z):
    # Complete this function to use the inputs x, y, and z to print the following:
    """
    The value of x is: <x>.
    The value of y is: <y>.
    The value of z is: <z>.
    """
    
    
    # YOUR CODE HERE
    print("The value of x is: {}.\nThe value of y is: {}.\nThe value of z is: {}.".format(x,y,z))

```

<!-- grade_id: cell-eedfb2417e94f554, points: 4, grade: true, solution: false -->
```python
from unittest.mock import patch
with patch('__main__.print') as mock_print:
    practice_2(1, 2, 3)
mock_print.assert_called_once_with('The value of x is: 1.\nThe value of y is: 2.\nThe value of z is: 3.')

with patch('__main__.print') as mock_print:
    practice_2(40, 50, 60)
mock_print.assert_called_once_with('The value of x is: 40.\nThe value of y is: 50.\nThe value of z is: 60.')

```

<!-- grade_id: cell-3f0ef5aae6375abf, points: 0, grade: false, solution: true -->
```python
# 4 points
def practice_3(length, width):
    # Complete this function to use the inputs length and width to 
    # calculate the area of a rectangle and pring:
    # The area of a rectangle with length <length> and width <width> is <area>.
    # Make sure to round the area to two decimal places. 
       
    # YOUR CODE HERE
    area = round(length * width,2)
    print("The area of a rectangle with length {} and width {} is {}.".format(length, width, area))
```

<!-- grade_id: cell-af0b79e6ea65b91b, points: 4, grade: true, solution: false -->
```python
from unittest.mock import patch
with patch('__main__.print') as mock_print:
    practice_3(45.23, 89.34)
mock_print.assert_called_once_with('The area of a rectangle with length 45.23 and width 89.34 is 4040.85.')

with patch('__main__.print') as mock_print:
    practice_3(40.90, 50.55)
mock_print.assert_called_once_with('The area of a rectangle with length 40.9 and width 50.55 is 2067.49.')

```

<!-- grade_id: cell-9fefb59228936ecb, points: 0, grade: false, solution: true -->
```python
# 4 points
def practice_4(radius):
    # Complete this function to return the surface area of a sphere given a radius. 
    # Use the input `radius` as the radius
    # round the surface area to two decimales 
    # Use the PI from the math module
    
    import math
    PI = math.pi
    
    # YOUR CODE HERE
    area = round(4*PI*radius**2,2)
    return area

```

<!-- grade_id: cell-491d00e6fd859523, points: 2, grade: true, solution: false -->
```python
# Test cell #1 for practice_4
assert practice_4(15) == 2827.43
assert practice_4(7.25) == 660.52
```

<!-- grade_id: cell-e7263d2ce679b3d9, points: 2, grade: true, solution: false -->
```python
# Test cell for practice_4 (hidden)
```

<!-- grade_id: cell-4bd7e48a470f11cb, points: 0, grade: false, solution: true -->
```python

def practice_5(radius):
    # Complete this function to return the volume of a sphere given a radius. 
    # Use the input `radius` as the radius
    # round the surface area to two decimales 
    # Use the PI from the math module
    
    import math
    PI = math.pi
    
    # YOUR CODE HERE
    volume = round((4/3)*PI*radius**3,2)
    return volume
```

<!-- grade_id: cell-eb0c9dab58f9c650, points: 2, grade: true, solution: false -->
```python
# Test cell #1 for practice_5
assert practice_5(15) == 14137.17
assert practice_5(7.5) == 1767.15
```

<!-- grade_id: cell-230780b7e56d019f, points: 2, grade: true, solution: false -->
```python
# Test cell for practice_5 (Hidden)
```

<!-- grade_id: cell-d72435be8ccbe507, points: 0, grade: false, solution: true -->
```python
def practice_6(weight, height):
    # The body mass index (BMI) is calculated as a peron's weight (in pounds) times 720, divided by the square
    # of the person's height (in inches). A BMI in the range 19-25, inclusive, is considered healthy. 
    # Complete this function to calculate and return a person's BMI 
    # Round the BMI to two decimals places
    
    # YOUR CODE HERE
    BMI = round((720*weight)/(height**2),2)
    return BMI
    
```

<!-- grade_id: cell-0a0e41eb9ea9b9d0, points: 2, grade: true, solution: false -->
```python
# Test cell  practice_6
assert practice_6(150, 65) == 25.56
```

<!-- grade_id: cell-f3b43da0aab8e5ce, points: 2, grade: true, solution: false -->
```python
# Test cell  for practice_6 (hidden)
```

<!-- grade_id: cell-1b4e4ef71c231d53, points: 0, grade: false, solution: true -->
```python
def practice_7(no_one_liter_bottles, no_more_than_one_liter_bottles):
    # When you buy soda bottles, you deposit a small amount to encourage recycling. 
    # For one liter bottle the depost is $0.10.
    # For a bottle larger than one liter, the depost is $0.25.
    # Complete this function to print a refund message like follows:
    # The refund is $<amount>.
    # Round the refund to two decimal places. 
    
    # YOUR CODE HERE
    refund = round(no_one_liter_bottles*0.10+no_more_than_one_liter_bottles*0.25,2)
    print('The refund amount is ${}.'.format(refund))
    
```

<!-- grade_id: cell-ee31a45ce013f115, points: 2, grade: true, solution: false -->
```python
# Test cell  practice_7
from unittest.mock import patch
with patch('__main__.print') as mock_print:
    practice_7(20, 50)
mock_print.assert_called_once_with('The refund amount is $14.5.')

```

<!-- grade_id: cell-1e4856019e1a79b4, points: 2, grade: true, solution: false -->
```python
# Test cell  for practice_7 (hidden)
```

<!-- grade_id: cell-3d3d9a85927e5aa7, points: 0, grade: false, solution: true -->
```python
def practice_8(a, b):
    # Complete this function to print the following given two numbers a and b
    """
    The sum of <a> and <b> is <result>.
    The difference when <b> is substracted from <a> is <result>.
    The product of <a> and <b> is <result>.
    The quotient when <a> is divided by <b> is <result>.
    The remainder when <a> is divided by <b> is <result>. 
    The result of <a>^<b> is <result>. 
    """
    
    """
    a = 10
    b = 5
    The sum of 10 and 5 is 15.
    The difference when 5 is substracted from 10 is -5.
    The product of 10 and 5 is 50.
    The quotient when 10 is divided by 5 is 2.0.
    The remainder when 10 is divided by 5 is 2.
    The result of 10^5 is 100000.

    
    """
    
    # YOUR CODE HERE
    
    sum = a + b
    sub = b - a
    product = a * b
    quotient = a/b
    remainder = a%b
    result = pow(a,b)
    print('The sum of {0} and {1} is {2}.\nThe difference when {0} is substracted from {1} is {3}.\nThe product of {0} and {1} is {4}.\nThe quotient when {0} is divided by {1} is {5}.\nThe remainder when {0} is divided by {1} is {6}.\nThe result of {0}^{1} is {7}.'.format(a,b,sum,sub,product,quotient,remainder,result))
```

<!-- grade_id: cell-ffc618cee68a002e, points: 2, grade: true, solution: false -->
```python
# Test cell  practice_8
from unittest.mock import patch
with patch('__main__.print') as mock_print:
    practice_8(10, 5)
result = """The sum of 10 and 5 is 15.
The difference when 10 is substracted from 5 is -5.
The product of 10 and 5 is 50.
The quotient when 10 is divided by 5 is 2.0.
The remainder when 10 is divided by 5 is 0.
The result of 10^5 is 100000."""

mock_print.assert_called_once_with(result)

```

<!-- grade_id: cell-faee29752b184c3a, points: 2, grade: true, solution: false -->
```python
# Test cell  for practice_8 (hidden)
```

<!-- grade_id: cell-1f5e00fa4aaca178, points: 0, grade: false, solution: true -->
```python
def practice_9(P, r, t):
    # The formula for simple interest is A = P(1+rt), where P is
    # the principle amount, r is the annual rate of interest,
    # t is the number of years the amount is invested, and A
    # is the amount at the end of the investment. 
    # Complete this function to print the following message:
    # After <t> years at <r>%, the investment will be worth $<A>.
    # Round the amount to two decimal places. 
    # Hint: rate is a percentage!
    
    # YOUR CODE HERE
    A = round(P*(1+r*t/100),2)
    print("After {} years at {}%, the investment will be worth ${}.".format(t,r,A))

```

<!-- grade_id: cell-5a5e0571e7a0c2bf, points: 2, grade: true, solution: false -->
```python
# Test cell  practice_9

from unittest.mock import patch
with patch('__main__.print') as mock_print:
    practice_9(1500, 4.3, 4)
mock_print.assert_called_once_with('After 4 years at 4.3%, the investment will be worth $1758.0.')

```

<!-- grade_id: cell-14b1955cc3c302df, points: 2, grade: true, solution: false -->
```python
# Test cell  for practice_9 (hidden)
```

<!-- grade_id: cell-5f5aaeb3dd44c165, points: 0, grade: false, solution: true -->
```python
def practice_10(P, r, t, n):
    # The formula for compound interest is 
    # A = P(1 + r/n)^nt
    # where:
    # P is the principle amount.
    # r is the annual rate of interest.
    # t is the number of years the amount is invested. 
    # n is the number of times the interest is compounded per year.
    # A is the amount at the end of the investment. 
    # Complete this function to print the following:
    # $<P> invested at <r>% for <t> years compounded <n> times per year is $<A>.

    
    # YOUR CODE HERE
     A = round(P*(1+r/(n*100))**(n*t),2)
     print("${} invested at {}% for {} years compounded {} times per year is ${}.".format(P,r,t,n,A))
    
    
```

<!-- grade_id: cell-30c68e3ca19ac5e1, points: 2, grade: true, solution: false -->
```python
# Test cell  practice_10

from unittest.mock import patch
with patch('__main__.print') as mock_print:
    practice_10(1500, 4.3, 6, 4)
mock_print.assert_called_once_with('$1500 invested at 4.3% for 6 years compounded 4 times per year is $1938.84.')

```

<!-- grade_id: cell-45ddeec89b139e2b, points: 2, grade: true, solution: false -->
```python
# Test cell  for practice_10 (hidden)
```