Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Vaishnavi Vukku"
COLLABORATORS = "Professor Zia"
```

---

<!-- grade_id: cell-ffc62b512fc653d1, points: 0, grade: false, solution: true -->
```python
# 2 points
def return_true():

    # return the value True
    return True
```

<!-- grade_id: cell-db4ecc495cb0c800, points: 2, grade: true, solution: false -->
```python
assert return_true() == True
```

<!-- grade_id: cell-c3f516fe563bd902, points: 0, grade: false, solution: true -->
```python
# 2 points
def return_false():

    # return the value False
    return False
```

<!-- grade_id: cell-be625703ceffd662, points: 2, grade: true, solution: false -->
```python
assert return_false() == False
```

<!-- grade_id: cell-72a401bac23e1f13, points: 0, grade: false, solution: true -->
```python
# 2 points
def return_int_10():
    # return 10
    return 10
```

<!-- grade_id: cell-c6f27b6b2157dfaf, points: 2, grade: true, solution: false -->
```python
assert return_int_10() == 10
```

<!-- grade_id: cell-064dabbc809caed3, points: 0, grade: false, solution: true -->
```python
# 2 points
def return_float_3_point_5():
    # return 3.5
    # YOUR CODE HERE
    return 3.5
```

<!-- grade_id: cell-fa8aca754b6c7261, points: 2, grade: true, solution: false -->
```python
assert return_float_3_point_5() == 3.5
```

<!-- grade_id: cell-25b3cc99451f7294, points: 0, grade: false, solution: true -->
```python
# 2 points
def return_letter_a():
    # return the letter a
    # YOUR CODE HERE
    return 'a'
```

<!-- grade_id: cell-44511ac54c8c7745, points: 2, grade: true, solution: false -->
```python
assert return_letter_a() == 'a'
```

<!-- grade_id: cell-af57e884088461ab, points: 0, grade: false, solution: true -->
```python
# 2 points
def return_string_EAS503():
    # return the string 'EAS503'
    # YOUR CODE HERE
    return 'EAS503'
```

<!-- grade_id: cell-cd85487f6b849359, points: 2, grade: true, solution: false -->
```python
assert return_string_EAS503() == 'EAS503'
```

<!-- grade_id: cell-b0eb174c1e273684, points: 0, grade: false, solution: true -->
```python
# 2 points
def exercise_1():
    # True or False: Buffalo is in New York.
    # Return True or False
    # YOUR CODE HERE
    return True

```

<!-- grade_id: cell-5c3de8d19d903ba1, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-b113c4c178b745c9, points: 0, grade: false, solution: true -->
```python
# 2 points
def exercise_2():

    # True or False: New York is the largest state in the United States
    # Return True or False
    # YOUR CODE HERE
    return False
```

<!-- grade_id: cell-4112ce1ca5f9ab79, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-70844679487f0ed9, points: 0, grade: false, solution: true -->
```python
# 2 points
def exercise_3():

    # Which of the following in an incorrect Python variable name? Return a letter.
    # (a) -- test1234
    # (b) -- 3pi
    # (c) -- __factor
    # (d) -- mean__value__2

    # YOUR CODE HERE
    return 'b'

```

<!-- grade_id: cell-e6b034753bf4cae5, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-b8915a0c6180c522, points: 0, grade: false, solution: true -->
```python
# 2 points
def exercise_4():

    # square 3 and 4 and return the their summed value
    # YOUR CODE HERE
    return (3**2)+(4**2) 
```

<!-- grade_id: cell-7ca0f8416918497b, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-dee7aa8efbfa5656, points: 0, grade: false, solution: true -->
```python
# 2 points
def exercise_5():

    # divide 2020 by 5 and return the value
    # YOUR CODE HERE
    return (2020/5)
```

<!-- grade_id: cell-1e2e40e16c6f331b, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-82ebd1a17cb3daa6, points: 0, grade: false, solution: true -->
```python
# 2 points
def exercise_6():

    # return the remainder of 512 divided by 16
    # YOUR CODE HERE
    return 512%16
```

<!-- grade_id: cell-c2c35898d5d590cf, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-31daed1e483a7192, points: 0, grade: false, solution: true -->
```python
# 2 points
def exercise_7():

    # return 36 divided by 5
    # YOUR CODE HERE
    return 36/5
```

<!-- grade_id: cell-36c85be2031c91d5, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-32aa3c2a01590949, points: 0, grade: false, solution: true -->
```python
# 2 points
def exercise_8():

    # return 36 divided by 5, BUT return only the quotient without the remainder
    # YOUR CODE HERE
    return 36//5
```

<!-- grade_id: cell-39b01facd75ef9d1, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-3df5e0dd68d4e331, points: 0, grade: false, solution: true -->
```python
# 2 points
def exercise_9():

    # return the area of a circle with radius 6
    # use the pi value given
    pi = 3.14
    # YOUR CODE HERE
    return pi*(6**2)
```

<!-- grade_id: cell-f023ed9880e092a5, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-07c06c3677239fa3, points: 0, grade: false, solution: true -->
```python
# 2 points
def exercise_10():

    # Given the temperature in fahrenheit (temperature_in_f), return the temperature in celcius
    temperature_in_f = 67
    # YOUR CODE HERE
    return (temperature_in_f - 32) * 5 / 9

```

<!-- grade_id: cell-515208d881ded972, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-de3bf51b813aea3d, points: 0, grade: false, solution: true -->
```python
# 2 points
def exercise_11():

    # Given the string 'EAS503', return the letter E using the square-bracket syntax

    class_name = 'EAS503'
    # YOUR CODE HERE
    return class_name[0]
```

<!-- grade_id: cell-11a0b4cf97edc78b, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-dd966809ddfe4f98, points: 0, grade: false, solution: true -->
```python
# 2 points
def exercise_12():

    # Given the string 'EAS503', return the letter S using the square-bracket syntax

    class_name = 'EAS503'
    # YOUR CODE HERE
    return class_name[2]
```

<!-- grade_id: cell-10cb53d5f5c19b9d, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-df92f9ba306b9349, points: 0, grade: false, solution: true -->
```python
# 2 points
def exercise_13():

    # Given the string 'EAS503', return the last character in the string using the square-bracket syntax

    class_name = 'EAS503'
    # YOUR CODE HERE
    return class_name[-1]
```

<!-- grade_id: cell-ceae1a5979d88643, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-dedafbc207653618, points: 0, grade: false, solution: true -->
```python
# 2 points
def exercise_14():

    # Given a float value, convert it to an int type and return the value

    a_float_value = 3.14159
    # YOUR CODE HERE
    return int(a_float_value)
```

<!-- grade_id: cell-81fdef9b15964916, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-f61bafc71e2e3a49, points: 0, grade: false, solution: true -->
```python
# 2 points
def exercise_15():

    # Pi is given in string data type, return it in float data type

    pi_string = '3.14159'
    # YOUR CODE HERE
    return float(pi_string)
```

<!-- grade_id: cell-45d4a5e676072c51, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-0d9832761fa4db55, points: 0, grade: false, solution: true -->
```python
# 2 points
def exercise_16():

    # Return the data type of the variable test_variable

    test_variable = '3.14159'
    # YOUR CODE HERE
    return type(test_variable)
```

<!-- grade_id: cell-504d1df50b3463ea, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-5e67f16c93a4fb62, points: 0, grade: false, solution: true -->
```python
# 2 points
def exercise_17():

    # Return the data type of the variable test_variable

    test_variable = 3.14159
    # YOUR CODE HERE
    return type(test_variable)
```

<!-- grade_id: cell-478224f7f29170f8, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-7d69a0f8377a7b37, points: 0, grade: false, solution: true -->
```python
# 2 points
def exercise_18():

    # True or False: In Python you can declare a variable without specifying its data type.
    # Return True or False
    # YOUR CODE HERE
    return True

```

<!-- grade_id: cell-7d785798f83aeb4d, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-e584cdd471a2f76a, points: 0, grade: false, solution: true -->
```python
# 2 points
def exercise_19():

    # True or False: In Python a variable that was used as a string can be reused as a float.
    # Return True or False
    # YOUR CODE HERE
    return True

```

<!-- grade_id: cell-3b39aec3d0af4b59, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-da3deb12139ffcd7, points: 0, grade: false, solution: true -->
```python
# 2 points
def exercise_20():

    # Given a height and base of a triangle, calculate its area and return it

    height = 3
    base = 5
    # YOUR CODE HERE
    return (base*height)/2
```

<!-- grade_id: cell-82de4e6f10a07f59, points: 2, grade: true, solution: false -->
```python

```