Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "VAISHNAVI VUKKU"
COLLABORATORS = "PROFESSOR ZIA"
```

---

<!-- grade_id: cell-cc604c0aae18cc24, points: 0, grade: false, solution: true -->
```python
# Exercise 1
def total_wages_for_the_week(hours, wage):
    # Many companies pay time-and-a-half for any hours worked above 40 in a given week.
    # Complete this function whose inputs are hours worked (hours) and the hourly rate (wage) to
    # calculate the total wages for the week. 
    # YOUR CODE HERE
    
    if hours > 40:
        extra_wage = (hours - 40)*(1.5*wage)
        total_wage = 40 * wage + extra_wage
    else:
        total_wage = hours * wage
    return total_wage
```

<!-- grade_id: cell-999bf8ee61b894e8, points: 2, grade: true, solution: false -->
```python
# Test cell for exercise 1
assert total_wages_for_the_week(40, 25) == 1000
assert total_wages_for_the_week(65, 63.4) == 4913.5
```

<!-- grade_id: cell-0c803c0254790b8a, points: 2, grade: true, solution: false -->
```python
# Test cell for exercise 1 (Hidden)
```

<!-- grade_id: cell-09452c74c9ab069a, points: 0, grade: false, solution: true -->
```python
# Exercise 2
def convert_quiz_score_to_letter_grade(score):
    # A certain CS professor gives five-point quizzes that are graded on the scale 5-A, 4-B, 3-C, 2-D, 1-F, 0-F.
    # Complete this function which accepts a quiz score as an input and uses a decision structure to calculate the corresponding
    # grade. 
    # YOUR CODE HERE
    if score == 0:
        grade = "F"
    elif score == 1:
        grade = "F"
    elif score == 2:
        grade = "D"
    elif score == 3:
        grade = "C"
    elif score == 4:
        grade = "B"
    else:
        grade = "A"
    return grade
```

<!-- grade_id: cell-212a6ed0aedc4322, points: 2, grade: true, solution: false -->
```python
# Test cell  for exercise 2 
assert convert_quiz_score_to_letter_grade(0) == "F"
assert convert_quiz_score_to_letter_grade(1) == "F"
assert convert_quiz_score_to_letter_grade(2) == "D"
assert convert_quiz_score_to_letter_grade(3) == "C"
assert convert_quiz_score_to_letter_grade(4) == "B"
assert convert_quiz_score_to_letter_grade(5) == "A"
```

<!-- grade_id: cell-c2699b0226be40b0, points: 0, grade: false, solution: true -->
```python
# Exercise 3 
def convert_exam_score_to_letter_grade(score):
    # A certain CS professor gives 100-point exams that are graded on the scale 90-100: A, 80-89: B, 70-79: C,
    # 69-69: D, < 60: F. Complete this function which accepts an exam score as input and uses a decision structure 
    # to calculate the corresponding grade. 
    # YOUR CODE HERE
    if 90<= score <= 100:
        grade = "A"
    elif 80 <= score <= 89:
        grade ="B"
    elif 70 <= score <= 79:
        grade ="C"
    elif 60 <= score <= 69:
        grade = "D"
    elif score < 60:
        grade = "F"
    else:
        print("please enter a valid score")
    
    return grade
```

<!-- grade_id: cell-7c59cc2534924674, points: 2, grade: true, solution: false -->
```python
# Test cell  for exercise 3
assert convert_exam_score_to_letter_grade(55) == "F"
assert convert_exam_score_to_letter_grade(35) == "F"
assert convert_exam_score_to_letter_grade(65) == "D"
assert convert_exam_score_to_letter_grade(73) == "C"
assert convert_exam_score_to_letter_grade(87) == "B"
assert convert_exam_score_to_letter_grade(92) == "A"
```

<!-- grade_id: cell-1507900e432c1220, points: 2, grade: true, solution: false -->
```python
# Test cell  for exercise 3(Hidden)
```

<!-- grade_id: cell-b9be67329b35ea84, points: 0, grade: false, solution: true -->
```python
# Exercise 4
def calculate_class_standing_from_credits(credits):
    # A certain college classifies students according to credits earned. A student with less than 7 credits is a Freshman.
    # At least 7 credits are required to be a Sophomore, 16 to be a Junior and 26 to be classifed as Senior. 
    # Complete this function which calculates the class standing from the number of credits earned. 
    # YOUR CODE HERE
    if credits < 7:
        student_classifier = "Freshman"
    elif 7 <= credits < 16:
        student_classifier = "Sophomore"
    elif 16 <= credits < 26:
        student_classifier = "Junior"
    elif credits >= 26:
        student_classifier = "Senior"
    return student_classifier
```

<!-- grade_id: cell-92d08237edebe22d, points: 2, grade: true, solution: false -->
```python
# Test cell  for exercise 4
assert calculate_class_standing_from_credits(6) == "Freshman"
assert calculate_class_standing_from_credits(13) == "Sophomore"
assert calculate_class_standing_from_credits(23) == "Junior"
assert calculate_class_standing_from_credits(45) == "Senior"
```

<!-- grade_id: cell-fc0a577e1c4515eb, points: 2, grade: true, solution: false -->
```python
# Test cell  for exercise 4
```

<!-- grade_id: cell-428c79d7f08e716a, points: 0, grade: false, solution: true -->
```python
# Exercise 5
def calculate_fine(limit, speed):
    # The speeding ticket fine policy in Podunksville is $50 plus $5 for each mph over the limit plus a 
    # penalty of $200 for any speed over 90 mph. Complete this function which accepts a speed limit and a clocked
    # speed and either return the string `Legal` or the amount of fine, if the speed is illegal. 

    # YOUR CODE HERE
    if speed <= 90 and speed > limit:
        extra_speed = speed - limit
        fine = 50 + (extra_speed * 5)
    elif speed > 90:
        extra_speed = speed - limit
        fine = 250 + (extra_speed * 5)
    else:
        fine = "Legal"
        
    return fine
```

<!-- grade_id: cell-ae9177af4cabd012, points: 3, grade: true, solution: false -->
```python
# Test cell  for exercise 5
assert calculate_fine(55, 85) == 200
assert calculate_fine(55, 45) == 'Legal'
assert calculate_fine(55, 100) == 475
```

<!-- grade_id: cell-7843b693298ffede, points: 3, grade: true, solution: false -->
```python
# Test cell  for exercise 5
```

<!-- grade_id: cell-6c3c0a14fe237c2e, points: 0, grade: false, solution: true -->
```python
# Exercise 6
def is_leap_year(year):
    # A year is a leap year if it is divisible by 4, unless it is a century year that is not divisible 
    # by 400. (1800 and 1900 are not leap years while 1600 and 2000 are). Complete this function which calculates
    # whether a year is a leap year. Return True or False

    # YOUR CODE HERE
    if (year % 4) == 0:
        if (year % 100) == 0:
            if (year % 400) ==0:
                leap_year = True
            else:
                leap_year = False
        else:
            leap_year = True
    else:
        leap_year = False
    
    return leap_year
```

<!-- grade_id: cell-2fac7294d0a4333c, points: 3, grade: true, solution: false -->
```python
# Test cell  for exercise 6
assert is_leap_year(2000) == True
assert is_leap_year(1800) == False
assert is_leap_year(1992) == True
```

<!-- grade_id: cell-d39435fe9fa2d644, points: 3, grade: true, solution: false -->
```python
# Test cell  for exercise 6
```

<!-- grade_id: cell-a76c5e81720d0e7b, points: 0, grade: false, solution: true -->
```python
# Exercise 7
def is_palindrome(input_string):
    # Completet this function which determines if the input_string is a palindrome. Return True or False
    # Use square brackets to reverse the input_string! Make sure to lower the input string before testing!
    
    # YOUR CODE HERE
    reversed_string = input_string[::-1]
    
    if input_string.lower() == reversed_string.lower():
        output = True
    else:
        output = False
    
    return output
```

<!-- grade_id: cell-cda9f94c36b5e077, points: 3, grade: true, solution: false -->
```python
# Test cell  for exercise 7
assert is_palindrome("Kayak") == True
assert is_palindrome("Rotator") == True
assert is_palindrome("AACA") == False

```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-7-c824b8d2a2e5> in <module>
      1 # Test cell  for exercise 7
----> 2 assert is_palindrome("Kayak") == True
      3 assert is_palindrome("Rotator") == True
      4 assert is_palindrome("AACA") == False

AssertionError:
```

<!-- grade_id: cell-049803af92bcfe13, points: 3, grade: true, solution: false -->
```python
# Test cell  for exercise 7 (hidden)
```

<!-- grade_id: cell-519e1e17ac49b564, points: 0, grade: false, solution: true -->
```python
# Ex8: Season from Month and Day
def determine_season(month, day):

    # The year is divided into four season: spring, summer, fall (or autumn) and winter.
    # While the exact dates that the seasons change vary a little bit from year to
    # year because of the way that the calender is constructed, we will use the following
    # dates for this exercise:

    # Season  -- First Day
    # Spring  -- March 20
    # Summer  -- June 21
    # Fall  -- September 22
    # Winter    -- December 21

    # Complete this function which takes in as inputs a month and day. It should
    # output the season.
    # input 1: month -- str
    # input 2: day -- int

    # output: month -- str (Spring, Summer, Fall, Winter)

    # YOUR CODE HERE
    
    month = str(month).capitalize()
    day = int(day)
    
    
    if month == "January" or month == "February":
        season = "Winter"
    elif month == "July" or month == "August":
        season = "Summer"
    elif month == "October" or month == "November":
        season = "Fall"
    elif month == "April" or month == "May":
        season = "Spring"
    else:
        if month == "March" and day in range(20,32):
            season = "Spring"
        elif month == "March" and day in range(1,20):
            season = "Winter"
        elif month == "June" and day in range(1,21):
            season = "Spring"
        elif month == "June" and day in range(21,30):
            season = "Summer"
        elif month == "September" and day in range(1,22):
            season = "Summer"
        elif month == "September" and day in range(22,31):
            season = "Fall"
        elif month == "December" and day in range(1,21):
            season = "Fall"
        elif month == "December" and day in range(21,32):
            season = "Winter"
        else:
            season = "Enter valid month"
    
    return season
```

<!-- grade_id: cell-4b8dd5ecc5f72de6, points: 2, grade: true, solution: false -->
```python
assert determine_season('March', 21) == 'Spring'
assert determine_season('June', 21) == 'Summer'
assert determine_season('November', 21) == 'Fall'
assert determine_season('January', 21) == 'Winter'
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-12-3c8e09dbfd21> in <module>
      1 assert determine_season('March', 21) == 'Spring'
----> 2 assert determine_season('June', 21) == 'Summer'
      3 assert determine_season('November', 21) == 'Fall'
      4 assert determine_season('January', 21) == 'Winter'

AssertionError:
```

<!-- grade_id: cell-deb5e4c4cab226f4, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-6f3941343261caca, points: 0, grade: false, solution: true -->
```python
#Ex9: Check a password

def check_password(password):

    # In this exercise you will complete this function to determine whether or not
    # a password is good. We will define a good password to be a one that is at least
    # 8 characters long and contains at least one uppercase letter, at least one lowercase
    # letter, and at least one number. This function should return True if the password
    # passed to it as its only parameter is good. Otherwise it should return False. 
    #
    # input: password (str)
    # output: True or False (bool)


    # YOUR CODE HERE
    upper_letter = False
    lower_letter = False
    numeric_value = False
    for i in password:
        if i.isupper():
            upper_letter = True
        elif i.islower():
            lower_letter = True
        elif i.isnumeric():
            numeric_value = True
        else:
            special_character = True
            
    if len(password)>=8 and upper_letter == True and lower_letter == True and numeric_value == True:
        password_good = True
    
    else:
        password_good = False
    
    return password_good

```

<!-- grade_id: cell-91c4927536aa6bdd, points: 2, grade: true, solution: false -->
```python
assert check_password('test1234') == False
assert check_password('password123') == False
assert check_password('SuperPasswrd90') == True
assert check_password('letmein!') == False

```

<!-- grade_id: cell-4eb94ff39a016bfe, points: 2, grade: true, solution: false -->
```python


```