Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "VAISHNAVI VUKKU"
COLLABORATORS = "PROFESSOR ZIA"
```

---

<!-- grade_id: cell-8abe173ecc588e38, points: 0, grade: false, solution: true -->
```python
# Ex1
def average_word_len_in_sentence(sentence):
    # Complete this function to calculate the average
    # word length in a sentence
    # Input: sentence
    # Output: average word length in sentence
    # Hint: count punctuations with whatever word they are `touching`
    # Hint: round the average to two decimal places
   
    # YOUR CODE HERE
    words = sentence.split()
    len_word = []
    sum_word = 0
    for i in range(len(words)):
        word = words[i]
        len_word.append(len(word))
    for i in range(len(len_word)):
        sum_word = sum_word + len_word[i]
    avg_word_ln = round(sum_word/len(len_word),2)
    return avg_word_ln
```

<!-- grade_id: cell-88eb75a67af14058, points: 2, grade: true, solution: false -->
```python
assert average_word_len_in_sentence('This is a test sentence!') == 4
```

<!-- grade_id: cell-9a48ceb083442619, points: 2, grade: true, solution: false -->
```python
# Hidden test
```

<!-- grade_id: cell-9cd6ed6844d07523, points: 0, grade: false, solution: true -->
```python
#Ex2
def wc(filename):
    # Complete this function to count the number of lines, words, and chars in a file. 
    # Input: filename
    # Output: a tuple with line count, word count, and char count -- in this order

    # YOUR CODE HERE
    num_lines = 0
    num_words = 0
    num_chars = 0
    with open(filename,'r') as file:
        for line in file:
            line_len = int(len(line))
            # print(line_len)
            line = line.strip()
            num_lines += 1
            words = line.split()
            num_words += len(words)
        
            num_chars += line_len
            
            
    return num_lines, num_words, num_chars

```

<!-- grade_id: cell-9784281c8327eef5, points: 4, grade: true, solution: false -->
```python
assert wc('ex2_data.txt') == (50, 643, 4298)

```

<!-- grade_id: cell-4cbfe85e128613b1, points: 0, grade: false, solution: true -->
```python
#Ex3a
def is_leap(year):
    # Complete this function to check if year is a leap year
    # Input: year
    # Output: True or False (Boolean)

    # YOUR CODE HERE
    leap = year % 4
    century = year % 100
    
    if leap == 0:
        if century == 0:
            if (year % 400) ==0:
                leap_year = True
            else:
                leap_year = False
        else:
            leap_year = True
    else:
        leap_year = False
    
    return leap_year


```

<!-- grade_id: cell-b9a05e2e268a2c6a, points: 4, grade: true, solution: false -->
```python
assert is_leap(1800) == False
assert is_leap(1900) == False
assert is_leap(1600) == True
assert is_leap(2000) == True
```

<!-- grade_id: cell-8b0875e64766bf5f, points: 0, grade: false, solution: true -->
```python
#Ex3b
def is_date_valid(month, day, year):
    # Complete this function to check if a data is valid, given month, day, and year.  
    # For example, 5/24/1962 is valid, but 9/31/2000 is not
    # Inputs: month, day, year
    # Output: True or False (Boolean)
    # Hint: Use is leap year function

    # YOUR CODE HERE
    if month == 2:
        if day == 29 and is_leap(year) == True:
            output = True
        elif day <= 28:
            output = True
        else:
            output = False
    elif month in (1, 3, 5, 7, 8, 10, 12):
        if day <= 31:
            output = True
        else:
            output = False
    elif month in (4, 6, 9, 11):
        if day <= 30:
            output = True
        else:
            output = False
    else:
        output = False
        
    return output
```

<!-- grade_id: cell-0af99a4a4a8d91ee, points: 2, grade: true, solution: false -->
```python
assert is_date_valid(5, 24, 1962) == True
assert is_date_valid(9, 31, 2000) == False
```

<!-- grade_id: cell-72635d6da9df7e2b, points: 2, grade: true, solution: false -->
```python
# Hidden tests
```

<!-- grade_id: cell-cdb696c1bc3474fa, points: 0, grade: false, solution: true -->
```python
#Ex4
def day_number(month, day, year):
    # Complete this function to calculate the day_number given month, day, and year.
    # Information: The days of the year are often numbered from 1 through 365 (or 366).
    # This number can be computed in three steps using int arithmetic:
    # (a) - day_num = 31 * (month - 1) + day
    # (b) - if the month is after February subtract (4*(month)+23)//10
    # (c) - if it's a leap year and after February 29, add 1
    # Hint: First verify the input date is valid, return False if it is not valid; use is_date_valid
    # Hint: Also use the is_leap function
    # Inputs: month, day, year
    # Output: the day number or False (boolean) if the date is invalid. 

    # YOUR CODE HERE
    if is_date_valid(month, day, year) == True:
        day_num = 31 * (month - 1) + day
        if month > 2 and is_leap(year) == False:
            day_num = day_num - ((4*(month)+23)//10)
        elif month > 2 and is_leap(year) == True:
            day_num = (day_num - (4*(month)+23)//10) + 1
        else:
            day_num = day_num
    else:
        day_num = False
    return day_num

```

<!-- grade_id: cell-411ea3b17d1adb33, points: 2, grade: true, solution: false -->
```python
assert day_number(9, 31, 2000) == False
assert day_number(2, 13, 2020) == 44
```

<!-- grade_id: cell-00925de994310b09, points: 2, grade: true, solution: false -->
```python
# Hidden tests

```

<!-- grade_id: cell-e03020ec7bc1b440, points: 0, grade: false, solution: true -->
```python
#Ex5
def years_to_double_investment(apr):
    # Complete this function to use a while loop to determine how long it takes for an investment 
    # to double at a given interest rate. The input to this function, apr, is the annualized interest rate
    # and the output is the number of years it takes an investment to double. Note: The amount of the initial 
    # investment (principal) does not matter; you can use $1. 
    # Hint: principal is the amount of money being invested. 
    # apr is the annual percentage rate expressed as a decimal number.  
    # Relationship: value after one year is given by principal * (1+ apr)

    # YOUR CODE HERE
    principal = 1
    year = 0
    bal = principal
    
    while bal < 2 * principal:
        bal = bal * (1+ apr)
        year += 1
    
    return year
```

<!-- grade_id: cell-37833043ebab338a, points: 2, grade: true, solution: false -->
```python
assert years_to_double_investment(0.06) == 12
assert years_to_double_investment(0.1) == 8
```

<!-- grade_id: cell-915174ac3c531b23, points: 0, grade: false, solution: true -->
```python
#Ex6 
def stopping_time(n):
    # Complete this function to return the number of steps taken to reach 1 in
    # the Collatz sequence (https://en.wikipedia.org/wiki/Collatz_conjecture)

    # YOUR CODE HERE
    step = 0
    while n != 1:
        if n % 2 == 0:
            n = n / 2
            step += 1
           
        else:
            n = 3*n + 1
            step += 1
            
    return step
```

<!-- grade_id: cell-cb2141fe7de88db3, points: 2, grade: true, solution: false -->
```python
assert stopping_time(12) == 9
assert stopping_time(27) == 111
```

<!-- grade_id: cell-a6ca4ec347d83750, points: 2, grade: true, solution: false -->
```python
# hidden tests
```

<!-- grade_id: cell-d5c69bc0cb4c6837, points: 0, grade: false, solution: true -->
```python
#Ex7
def is_prime(n):
    # A positive whole number > 2 is prime if no number between 2 and sqrt(n)
    # (include) evenly divides n. Write a program that accepts a value of n as
    # input and determine if the value is prime. If n is not prime, your program should
    # return False (boolean) as soon as it finds a value that evenly divides n.
    # Hint: if number is 2, return False

    import math
    
    # YOUR CODE HERE
    output = True
    if n ==2:
        output = False
    for i in range(2,int( math.sqrt(n)+1)):
        if n % i == 0:
            output = False
            break
        else:
            output = True
            
    return output
```

<!-- grade_id: cell-da18c35e0a5d579e, points: 4, grade: true, solution: false -->
```python
assert is_prime(2) == False
assert is_prime(3) == True
assert is_prime(5) == True
assert is_prime(25) == False

```

<!-- grade_id: cell-ad04b1135365e292, points: 0, grade: false, solution: true -->
```python
#Ex8
def all_primes(n):
    # Complete this function to return all the primes as a list less than or equal to n
    # Input: n
    # Output: a list of numbers

    # YOUR CODE HERE
    all_primes = [3]
    for primes in range(4, n+1):
        if is_prime(primes) == True:
            all_primes.append(primes)
        else:
            continue
                
    return all_primes
```

<!-- grade_id: cell-3e7dbc25ba03421c, points: 2, grade: true, solution: false -->
```python
assert all_primes(5) == [3, 5]
assert all_primes(25) == [3, 5, 7, 11, 13, 17, 19, 23]

```

<!-- grade_id: cell-d6cdbdb6ba20e916, points: 0, grade: false, solution: true -->
```python
#Ex9
def gcd(m,n):
    # Complete this function to determine the greatest common divisor (GCD). 
    # The GCD of two values can be computed using Euclid's algorithm. Starting with the values
    # m and n, we repeatedly apply the formula: n, m = m, n%m until m is 0. At this point, n is the GCD
    # of the original m and n. 
    # Inputs: m and n which are both natural numbers
    # Output: gcd

    # YOUR CODE HERE
    if m < n:
        a = n
        b = m
    else:
        a = m
        b = n
    
    while a != 0:
        r = a%b
        
        if r == 0:
            break
        else:
            a = b
            b = r
    return b

```

<!-- grade_id: cell-9a4f50a09e28d52d, points: 2, grade: true, solution: false -->
```python
assert gcd(25,75) == 25
assert gcd(3,13) == 1
```

<!-- grade_id: cell-6bba97d2c8fa2d9f, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-2e1d0f8cbd0fefe2, points: 0, grade: false, solution: true -->
```python
#Ex10
def determine_min_max_average(filename):
    # Complete this function to read grades from a file and determine the student with the highest average 
    # test grades and the lowest average test grades. 
    # Input: filename
    # Output: a tuple containing four elements: name of student with highest average, their average, 
    # name of the student with the lowest test grade, and their average. Example ('Student1', 99.50, 'Student5', 65.50)
    # Hint: Round to two decimal places

    # YOUR CODE HERE
    student_names = []
    average_scores = []
    # combined_list = []

    with open(filename,'r') as file:
        for line in file:
            line = line.strip()
            line = list(line.split(","))
            
            if len(line) == 26:
                subjects = len(line)-1
                sum = 0
        
                for i in range(len(line)):
                    if i == 0:
                        student_name = line[i]
                        student_names.append(student_name)
                    else:
                        sum = sum + int(line[i])
                average = round(sum/subjects,2)
                average_scores.append(average)
            else:
                continue
        
        max_index = average_scores.index(max(average_scores))
        max_student_name = student_names[max_index]
        max_average = max(average_scores)
        min_index = average_scores.index(min(average_scores))
        min_student_name = student_names[min_index]
        min_average = min(average_scores)
   
    return max_student_name, max_average, min_student_name, min_average
```

<!-- grade_id: cell-c8ec04f59adc84af, points: 4, grade: true, solution: false -->
```python
assert determine_min_max_average('ex10_data.txt') == ('student733', 86.2, 'student202', 65.4)
```

<!-- grade_id: cell-423f49b4dbb9cda6, points: 0, grade: false, solution: true -->
```python
# Ex11:Is a License Plate Valid?

def is_license_plate_valid(plate):

    # In a particular jurisdiction, older license plates consist of three uppercase
    # letters followed by three digits. When all of the license plates following
    # that pattern had been used, the format was changed to four digits followed by
    # three uppercase letters. 

    # Complete this function whose only input is a license plate and its output
    # is: 1) Older/Valid 2) Newer/Valid 3) Invalid
    # input: plate (str)
    # output: 'Older/Valid' or 'Newer/Valid' or 'Invalid'

    # YOUR CODE HERE
    plate_lst = list(plate)
    if len(plate_lst) == 6 or len(plate_lst) == 7:
        if plate_lst[0].isupper() and plate_lst[1].isupper() and plate_lst[2].isupper() and plate_lst[3].isnumeric() and plate_lst[4].isnumeric() and plate_lst[5].isnumeric():
            output = 'Older/Valid'
        elif plate_lst[0].isnumeric() and plate_lst[1].isnumeric() and plate_lst[2].isnumeric() and plate_lst[3].isnumeric() and plate_lst[4].isupper() and plate_lst[5].isupper() and plate_lst[6].isupper():
            output = 'Newer/Valid' 
        else:
            output = 'Invalid'
    else:
        output = 'Invalid'
    return output

```

<!-- grade_id: cell-1900e84b3f513c64, points: 4, grade: true, solution: false -->
```python
assert is_license_plate_valid('ABC123') == 'Older/Valid'
assert is_license_plate_valid('GHE952') == 'Older/Valid'
assert is_license_plate_valid('1934ABT') == 'Newer/Valid'
assert is_license_plate_valid('bTR342') == 'Invalid'
```

<!-- grade_id: cell-cfcb59d207e5d3cd, points: 4, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-859f437c63175722, points: 0, grade: false, solution: true -->
```python
# Ex12: Magic dates

def is_magic_date(date):
    # A magic date is a date where the day multiplied by the month is equal 
    # to the two digit year. For example, June 10, 1960 is a magic date because
    # June is the sixth month, and 6 times 10 is 60, which is equal to the two
    # digit year. Complete this function to determine whether or not a date is
    # a magic date.

    # input: date (str -- month/day/year)
    # output: True or False (bool)


    # YOUR CODE HERE
    date_lst = list(date.split("/")) 
    
    day_month = int(date_lst[0]) * int(date_lst[1])
    year = date_lst[2]
    year = year[2:]
    
    if int(day_month) == int(year):
        output = True
    else:
        output = False
    
    return output
```

<!-- grade_id: cell-7a1fa98499aba445, points: 2, grade: true, solution: false -->
```python
assert is_magic_date('6/10/1960') == True
assert is_magic_date('6/8/1948') == True
```

<!-- grade_id: cell-22929b608ed4e3ba, points: 0, grade: false, solution: true -->
```python
# Ex13
def remove_outliers(data, num_outliers):
    # When analyzing data collected as a part of a science experiment it 
    # may be desriable to remove the most extreme values before performing
    # other calculations. Complete this function which takes a list of
    # values and an non-negative integer, num_outliers, as its parameters.
    # The function should create a new copy of the list with the num_outliers 
    # largest elements and the num_outliers smallest elements removed. 
    # Then it should return teh new copy of the list as the function's only 
    # result. The order of the elements in the returned list does not have to
    # match the order of the lemetns in the original list.
    # input1: data (list)
    # input2: num_outliers (int)

    # output: list



    # YOUR CODE HERE
    pre_process = data.sort()
    pre_process = data[num_outliers:-num_outliers]  
    
       
    return pre_process
```

<!-- grade_id: cell-f37107e0e303a5ca, points: 2, grade: true, solution: false -->
```python
import random
random.seed(1234)
data = [random.randint(50, 150) for ele in range(100)]
data[45] = 1
data[46] = 2
data[90] = 250
data[34] = 300

result = [50, 50, 51, 52, 52, 52, 53, 53, 54, 
55, 55, 55, 55, 56, 58, 58, 59, 59, 59, 60, 61, 61, 
61, 62, 64, 64, 64, 68, 68, 68, 69, 69, 69, 70, 71, 
72, 73, 75, 77, 80, 81, 81, 84, 84, 84, 85, 88, 88, 89, 
92, 94, 94, 95, 95, 98, 103, 106, 108, 108, 109, 109, 109, 
110, 111, 111, 112, 113, 114, 115, 115, 117, 117, 119, 121, 
124, 124, 125, 126, 128, 129, 132, 132, 133, 134, 135, 135, 
135, 136, 138, 140, 141, 148, 148, 149, 149, 150]

assert remove_outliers(data, 2) == result

```

<!-- grade_id: cell-f584e176be7f53bc, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-76c8326115a8aee3, points: 0, grade: false, solution: true -->
```python
# Ex14: Removing duplicates

def remove_duplicates(words):
    # Complete this function to remove duplicates from the words list using a loop
    # input: words (list)
    # output: a list without duplicates

    # YOUR CODE HERE
    output_list = []
    
    for i in words:
        
        if i not in output_list:
            output_list.append(i)
    
    return output_list


```

<!-- grade_id: cell-0b3ad78a5f56ef26, points: 2, grade: true, solution: false -->
```python
assert remove_duplicates([1,2,3,3,3,4,5,6,7,7,8]) == [1, 2, 3, 4, 5, 6, 7, 8]

```

<!-- grade_id: cell-6f1773c6eb139b8d, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-3b7cdbf9f92b6fb9, points: 0, grade: false, solution: true -->
```python
# Ex15: List of proper divisors

def proper_divisors(n):
    # A proper divisor ofa  positive integer, n, is a positive integer less than n which divides
    # evenly into n. Complete this function to compute all the proper divisors of a positive
    # integer. The integer is passed to this function as the only parameter. The function will
    # return a list of containing all of the proper divisors as its only reult. 

    # input: n (int)
    # output: list

    # YOUR CODE HERE
    div_lst = []
    
    for i in range(n-1):
        i = i+1
        if n%i == 0:
            div_lst.append(i)
        else:
            continue
    
    return div_lst

```

<!-- grade_id: cell-13e6b6403196da76, points: 2, grade: true, solution: false -->
```python
assert proper_divisors(28) == [1, 2, 4, 7, 14]

```

<!-- grade_id: cell-e7223c2aeb9c5d36, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-a94951aac16e950f, points: 0, grade: false, solution: true -->
```python
#Ex16: Perfect Numbers
def is_number_perfect(n):
    # An integer, n, is said to be perfect when the sum of all of the proper divisors 
    # of n is equal to n. For example, 28 is a perfect number because its proper divisors
    # are 1, 2, 4, 7, and 14 = 28 
    # Complete this function to determine if a the number a perfect number or not. 
    # input: n (int)
    # output: True or False (bool)


    # YOUR CODE HERE
    lst_div = proper_divisors(n)
    print(lst_div)
    sum = 0
    for i in range(len(lst_div)):
        sum = sum + lst_div[i]
      
    if sum == n:
        output = True
    else:
        output = False
    
    return output
```

<!-- grade_id: cell-14d8cfff6ced5461, points: 2, grade: true, solution: false -->
```python
assert is_number_perfect(28) == True
assert is_number_perfect(76) == False

```

Output:
```
[1, 2, 4, 7, 14]
[1, 2, 4, 19, 38]
```

<!-- grade_id: cell-14eb44447e299251, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-21d1d1b24e27c644, points: 0, grade: false, solution: true -->
```python
# Ex17: Line of Best Fit
def best_line(points):
    # Complete this function to determine the best line. 
    # https://www.varsitytutors.com/hotmath/hotmath_help/topics/line-of-best-fit
    # input: points (list of tuples contain x, y values)
    # output: (m, b) # round both values to two decimal places

    # YOUR CODE HERE
     #step1
    x_list = [x_crd[0] for x_crd in points]
    x_sum = 0
    for i in range(len(x_list)):
        x_sum = x_sum + x_list[i]
    mean_x = x_sum/len(x_list)
    
    y_list = [y_crd[1] for y_crd in points]
    y_sum = 0
    for j in range(len(y_list)):
        y_sum = y_sum + y_list[j]
    mean_y = y_sum/len(y_list)
    
    #step2
    num = 0
    den = 0
    for i in range(len(x_list)):
        num = num + ((x_list[i] - mean_x)*(y_list[i] - mean_y))
        den = den + ((x_list[i] - mean_x)**2)
    slope = round(num/den,2)
    
    #step3
    y_intercept = round(mean_y - (num/den * mean_x),2)
                    
    return slope, y_intercept

```

<!-- grade_id: cell-386bd6a9301a19da, points: 2, grade: true, solution: false -->
```python
points = (
    (8, 3),
    (2, 10),
    (11, 3),
    (6, 6),
    (5, 8),
    (4, 12),
    (12, 1),
    (9, 4),
    (6, 9),
    (1, 14),
)

assert best_line(points) == (-1.11, 14.08)
```

<!-- grade_id: cell-e3b40327cdd21344, points: 2, grade: true, solution: false -->
```python

```

```python

```